# Statystyka matematyczna - ćwiczenia laboratoryjne

Uwagi i wytyczne ogólne dotyczące uzupełniania notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

---

# Zestaw zadań 8: Estymacja metodą bootstrap

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Dane do zadań

W celu wygenerowania danych wykorzystywanych w zawartych w notatniku zadaniach i komórkach testowych wykonaj obie poniższe komórki.

```python
sample_1 = pd.DataFrame(data=sp.stats.norm.rvs(loc=5, scale=0.2, size=20, random_state=7), columns=["X"])
sample_1.head()
```

Output:
```
X
0  5.338105
1  4.906813
2  5.006564
3  5.081503
4  4.842215
```

```python
x = sp.stats.uniform.rvs(loc=-2, scale=10, size=25, random_state=34)
y = 2*x - 5 + sp.stats.norm.rvs(loc=0, scale=2, size=25, random_state=13)
sample_2 = pd.DataFrame(data=np.array([x, y]).T, columns=["X", "Y"])
sample_2.head()
```

Output:
```
X         Y
0 -1.614383 -9.653548
1  5.801005  8.109542
2 -1.072962 -7.234931
3  4.328927  4.561478
4 -1.861092 -6.031980
```

### Zadanie 1: Estymacja bootstapowa parametrów próby losowej [8 pkt]

Przygotuj funkcję `bootstrap_parameter_estimation()`, która korzystając z metody bootstrap będzie dokonywała estymacji przedziałowej podstawowych parametrów próby losowej - wartości oczekiwanej, wariancji i odchylenia standardowego dla danych wejściowych w postaci szeregu szczegółowego zgodnie ze schematem z załączonego do notatnika zestawu wzorów.

W obliczeniach skorzystaj z nieobciążonych estymatorów wariancji i odchylenia standardowego.

Oprócz zwracanych wartości granic przedziału funkcja powinna wyświetlać następujący komunikat:

`P należy do przedziału [X, Y] przy założeniu poziomu ufności 1-𝛼 = Z`

gdzie P (symbol lub nazwa estymowanego parametru), X, Y i Z są automatycznie uzupełniane przez funkcję, a X i Y dodatkowo sformatowane w taki sposób, żeby wyświetlały się z dokładnością do 4 miejsc po przecinku.

Uwagi do wykonania zadania:
 - Automatyczne testy zakładają losowanie prób wg następującego schematu - kolejne próby są losowane w pęli for, a funkcja losująca próbę (np. `pd.sample()`) w kolejnych losowaniach ma ustawioną wartość parametru odpowiadającego za ziarno generatora liczb losowych na wartość `random_state`, `random_state+1`, `random_state+2` itd. 
 - Celem zadania jest napisanie funkcji od podstaw, w rozwiązaniu nie należy korzystać z gotowych funkcji służących do estymacji bootstrapowej.

<!-- grade_id: cell-618e12c8ffef0965, points: 0, grade: false, solution: true -->
```python
def bootstrap_parameter_estimation(data, parameter, number_of_samples, alpha, random_state):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks i pojedynczą kolumnę "X" z wartościami próby losowej.
    parameter: str
        W zależności od estymowanego parametru przyjmuje wartość "mean", "var" lub "std".
    number_of_samples: int
        Liczba prób bootstrapowych.
    alpha: float
        Wartość określająca poziom istotności.
    random_state: int
        Ziarno generatora liczb pseudolosowych.
    
    Returns
    -------
    parameter_low: float
        Dolna granica wyliczonego przedziału ufności.
    parameter_high: float
        Górna granica wyliczonego przedziału ufności.
    """       
    # YOUR CODE HERE
    x = data['X'].values
    n = len(x)
    estimates = []

    for i in range(number_of_samples):
        sample = np.random.RandomState(seed=random_state + i).choice(x, size=n, replace=True)

        if parameter == "mean":
            est = np.mean(sample)
        elif parameter == "var":
            est = np.var(sample, ddof=1)  
        elif parameter == "std":
            est = np.std(sample, ddof=1) 

        estimates.append(est)

    lower = np.percentile(estimates, 100 * alpha / 2)
    upper = np.percentile(estimates, 100 * (1 - alpha / 2))
    Z= 1 - alpha

    P = {
        "mean": "Wartość oczekiwana",
        "var": "Wariancja",
        "std": "Odchylenie standardowe"
    }[parameter]

    print(f"{P} należy do przedziału [{lower:.4f}, {upper:.4f}] przy założeniu poziomu ufności 1-𝛼 = {Z:.4f}")

    return lower, upper
```

<!-- grade_id: cell-eab0d556900c97d4, points: 8, grade: true, solution: false -->
```python
# Komórka testowa
assert np.all(np.isclose(bootstrap_parameter_estimation(sample_1, "mean", 1000, 0.05, 34), (4.913578501079754, 5.058598650238277)))
assert np.all(np.isclose(bootstrap_parameter_estimation(sample_1, "var", 2000, 0.02, 31), (0.009780029754528753, 0.05012098474121563)))
assert np.all(np.isclose(bootstrap_parameter_estimation(sample_1, "std", 1500, 0.1, 39), (0.11848597559111698, 0.20690074715252046)))
```

Output:
```
Wartość oczekiwana należy do przedziału [4.9136, 5.0586] przy założeniu poziomu ufności 1-𝛼 = 0.9500
Wariancja należy do przedziału [0.0098, 0.0501] przy założeniu poziomu ufności 1-𝛼 = 0.9800
Odchylenie standardowe należy do przedziału [0.1185, 0.2069] przy założeniu poziomu ufności 1-𝛼 = 0.9000
```

### Zadanie 2: Estymacja bootstrapowa wartości współczynnika korelacji liniowej [4 pkt]

Przygotuj funkcję `bootstrap_correlation_estimation()`, która korzystając z metody bootstrap będzie dokonywała estymacji przedziałowej wartości współczynnika korelacji liniowej Pearsona dla danych wejściowych w postaci szeregu szczegółowego zgodnie ze schematem z załączonego do notatnika zestawu wzorów.

Oprócz zwracanych wartości granic przedziału funkcja powinna wyświetlać następujący komunikat:

`r należy do przedziału [X, Y] przy założeniu poziomu ufności 1-𝛼 = Z`

gdzie X, Y i Z są automatycznie uzupełniane przez funkcję, a X i Y dodatkowo sformatowane w taki sposób, żeby wyświetlały się z dokładnością do 4 miejsc po przecinku.

Uwagi do wykonania zadania:
 - Automatyczne testy zakładają losowanie prób wg następującego schematu - kolejne próby są losowane w pęli for, a funkcja losująca próbę (np. `pd.sample()`) w kolejnych losowaniach ma ustawioną wartość parametru odpowiadającego za ziarno generatora liczb losowych na wartość `random_state`, `random_state+1`, `random_state+2` itd. 
 - Celem zadania jest napisanie funkcji od podstaw, w rozwiązaniu nie należy korzystać z gotowych funkcji służących do estymacji bootstrapowej.

<!-- grade_id: cell-d785d7ff406d589d, points: 0, grade: false, solution: true -->
```python
def bootstrap_correlation_estimation(data, number_of_samples, alpha, random_state):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks i dwie kolumny "X" i "Y" z wynikami próby losowej.
    number_of_samples: int
        Liczba prób bootstrapowych.
    alpha: float
        Wartość określająca poziom istotności.
    random_state: int
        Ziarno generatora liczb pseudolosowych.
    
    Returns
    -------
    r_corr_low: float
        Dolna granica wyliczonego przedziału ufności.
    r_corr_high: float
        Górna granica wyliczonego przedziału ufności.
    """
    # YOUR CODE HERE
    n = len(data)
    r_boot = []
    for i in range(number_of_samples):
        samp = data.sample(n=n, replace=True, random_state=random_state + i)
        r_boot.append(samp["X"].corr(samp["Y"]))
    r_boot = np.array(r_boot)
    lower = np.percentile(r_boot, 100 * (alpha / 2))
    upper = np.percentile(r_boot, 100 * (1 - alpha / 2))
    conf = 1 - alpha
    print(f"r należy do przedziału [{lower:.4f}, {upper:.4f}] przy założeniu poziomu ufności 1-𝛼 = {conf}")
    return lower, upper
```

<!-- grade_id: cell-4dfca86f32f1359a, points: 4, grade: true, solution: false -->
```python
# Komórka testowa
assert np.all(np.isclose(bootstrap_correlation_estimation(sample_2, 1200, 0.05, 54), (0.9292256084679922, 0.9816422271248535)))
assert np.all(np.isclose(bootstrap_correlation_estimation(sample_2, 1500, 0.02, 14), (0.9171033819425604, 0.9840550768251518)))
```

Output:
```
r należy do przedziału [0.9292, 0.9816] przy założeniu poziomu ufności 1-𝛼 = 0.95
r należy do przedziału [0.9171, 0.9841] przy założeniu poziomu ufności 1-𝛼 = 0.98
```

### Zadanie 3: Estymacja bootstrapowa współczynników równania regresji liniowej [6 pkt]

Przygotuj funkcję `bootstrap_linear_regression_coefficients_estimation()`, która korzystając z metody bootstrap będzie dokonywała estymacji przedziałowej wartości współczynników równania regresji liniowej dla danych wejściowych w postaci szeregu szczegółowego zgodnie ze schematem z załączonego do notatnika zestawu wzorów.

Oprócz zwracanych wartości granic przedziału funkcja powinna wyświetlać następujący komunikat:

`a należy do przedziału [X1, Y1] przy założeniu poziomu ufności 1-𝛼 = Z1` </br>
`b należy do przedziału [X2, Y2] przy założeniu poziomu ufności 1-𝛼 = Z2`

gdzie X1, X2, Y1, Y2, Z1 i Z2 są automatycznie uzupełniane przez funkcję, a X1, X2, Y1 i Y2 dodatkowo sformatowane w taki sposób, żeby wyświetlały się z dokładnością do 4 miejsc po przecinku.

Uwagi do wykonania zadania:
 - Automatyczne testy zakładają losowanie prób wg następującego schematu - kolejne próby są losowane w pęli for, a funkcja losująca próbę (np. `pd.sample()`) w kolejnych losowaniach ma ustawioną wartość parametru odpowiadającego za ziarno generatora liczb losowych na wartość `random_state`, `random_state+1`, `random_state+2` itd. 
 - Celem zadania jest napisanie funkcji od podstaw, w rozwiązaniu nie należy korzystać z gotowych funkcji służących do estymacji bootstrapowej.

<!-- grade_id: cell-03ff8749dfe56b09, points: 0, grade: false, solution: true -->
```python
def bootstrap_linear_regression_coefficients_estimation(data, number_of_samples, alpha, random_state):
    """
    Parameters
    -------
    data: DataFrame
        Tablica zawierająca domyślny indeks i dwie kolumny "X" i "Y" z wynikami próby losowej.
    number_of_samples: int
        Liczba prób bootstrapowych.
    alpha: float
        Wartość określająca poziom istotności.
    random_state: int
        Ziarno generatora liczb pseudolosowych.
    
    Returns
    -------
    a_ci: tuple
        Zmienna typu tuple zawierajaca granice przedziału ufności parametru a (a_low, a_high).
    b_ci: tuple
        Zmienna typu tuple zawierajaca granice przedziału ufności parametru b (b_low, b_high).
    """            
    # YOUR CODE HERE
    n = len(data)
    slopes = []
    intercepts = []

    # Bootstrap
    for i in range(number_of_samples):
        samp = data.sample(n=n, replace=True, random_state=random_state + i)
        x = samp["X"]
        y = samp["Y"]
        x_mean = x.mean()
        y_mean = y.mean()
        # nieobciążony estymator kowariancji
        cov_xy = ((x - x_mean) * (y - y_mean)).sum() / (n - 1)
        var_x = x.var(ddof=1)
        a = cov_xy / var_x
        b = y_mean - a * x_mean

        slopes.append(a)
        intercepts.append(b)

    slopes = np.array(slopes)
    intercepts = np.array(intercepts)

    # Percentylowy przedział ufności
    lower_q = 100 * (alpha / 2)
    upper_q = 100 * (1 - alpha / 2)

    a_low, a_high = np.percentile(slopes, [lower_q, upper_q])
    b_low, b_high = np.percentile(intercepts, [lower_q, upper_q])
    conf = 1 - alpha

    print(f"a należy do przedziału [{a_low:.4f}, {a_high:.4f}] przy założeniu poziomu ufności 1-𝛼 = {conf}")
    print(f"b należy do przedziału [{b_low:.4f}, {b_high:.4f}] przy założeniu poziomu ufności 1-𝛼 = {conf}")

    return (a_low, a_high), (b_low, b_high)
```

<!-- grade_id: cell-0febc4ee671ad799, points: 6, grade: true, solution: false -->
```python
# Komórka testowa
assert np.all(np.isclose(bootstrap_linear_regression_coefficients_estimation(sample_2, 1800, 0.05, 71), ((1.8361393329086446, 2.2918023790827124), (-5.575502208944547, -3.6854174758043574))))
print()
assert np.all(np.isclose(bootstrap_linear_regression_coefficients_estimation(sample_2, 2100, 0.02, 62), ((1.7964929430496048, 2.337236979144576), (-5.771238137193083, -3.481203748551042))))
```

Output:
```
a należy do przedziału [1.8361, 2.2918] przy założeniu poziomu ufności 1-𝛼 = 0.95
b należy do przedziału [-5.5755, -3.6854] przy założeniu poziomu ufności 1-𝛼 = 0.95

a należy do przedziału [1.7965, 2.3372] przy założeniu poziomu ufności 1-𝛼 = 0.98
b należy do przedziału [-5.7712, -3.4812] przy założeniu poziomu ufności 1-𝛼 = 0.98
```

### Zadanie 4: Estymacja bootstrapowa za pomocą dedykowanych funkcji [7 pkt]

Korzystając z funkcji `scipy.stats.bootstrap()` dokonaj estymacji przedziałowej następujących parametrów na poziomie istotności 1-alpha = 0.98:
 - na podstawie zmiennej `sample_1`: średnia, wariancja, odchylenie standardowe,
 - na podstawie zmiannej `sample_2`: współczynnik korelacji liniowej Pearsona, współczynniki równania regresji liniowej a i b.
 
Wyniki zapisz do zmiennych o nazwach `mean_ci`, `var_ci`, `std_ci`, `r_ci`, `a_ci` oraz `b_ci` zawierających listy z dolnymi i górnymi granicami przedziałów ufności. Podczas obliczania wartości każdego z parametrów ustaw wartość argumentu funkcji `scipy.stats.bootstrap()` odpowiadającego za ziarno generatora liczb losowych na 25.

Uwagi do wykonania zadania:
 - Automatyczne testy przewidują obliczenie wartości wariancji i odchylenia standardowego dla wartości parametru `ddof=0`.

<!-- grade_id: cell-5028ed987941ac51, points: 0, grade: false, solution: true -->
```python
# mean
import numpy as np
from scipy.stats import bootstrap

alpha = 0.02
conf_level = 1 - alpha
seed = 25

# 1) mean_ci dla sample_1
mean_res = bootstrap(
    data=(sample_1["X"].values,),
    statistic=np.mean,
    confidence_level=conf_level,
    random_state=seed
)
mean_ci = [mean_res.confidence_interval.low,
           mean_res.confidence_interval.high]


```

<!-- grade_id: cell-1b6eb85e5b1da27a, points: 0, grade: false, solution: true -->
```python
# var
import numpy as np
from scipy.stats import bootstrap

alpha = 0.02
conf_level = 1 - alpha
seed = 25

# 2) var_ci dla sample_1 (ddof=0)
def var_stat(x, axis):
    return np.var(x, axis=axis, ddof=0)

var_res = bootstrap(
    data=(sample_1["X"].values,),
    statistic=var_stat,
    confidence_level=conf_level,
    random_state=seed
)
var_ci = [var_res.confidence_interval.low,
          var_res.confidence_interval.high]


```

<!-- grade_id: cell-7fa6f2c46b14d185, points: 0, grade: false, solution: true -->
```python
# std
import numpy as np
from scipy.stats import bootstrap

alpha = 0.02
conf_level = 1 - alpha
seed = 25

# 3) std_ci dla sample_1 (ddof=0)
def std_stat(x, axis):
    return np.std(x, axis=axis, ddof=0)

std_res = bootstrap(
    data=(sample_1["X"].values,),
    statistic=std_stat,
    confidence_level=conf_level,
    random_state=seed
)
std_ci = [std_res.confidence_interval.low,
          std_res.confidence_interval.high]
```

<!-- grade_id: cell-185444e7b6b077c2, points: 0, grade: false, solution: true -->
```python
# r
import numpy as np
from scipy.stats import bootstrap

alpha = 0.02
conf_level = 1 - alpha
seed = 25

# 4) r_ci dla sample_2 – korelacja Pearsona
def corr_stat(x, y, axis=0):
    return np.corrcoef(x, y)[0, 1]

r_res = bootstrap(
    data=(sample_2["X"].values, sample_2["Y"].values),
    statistic=corr_stat,
    paired=True,
    vectorized=False,
    confidence_level=conf_level,
    random_state=seed
)
r_ci = [r_res.confidence_interval.low,
        r_res.confidence_interval.high]


```

<!-- grade_id: cell-6c7a615fa2793828, points: 0, grade: false, solution: true -->
```python
# a
import numpy as np
from scipy.stats import bootstrap

alpha = 0.02
conf_level = 1 - alpha
seed = 25

# 5) a_ci – nachylenie prostej regresji (populacyjna kowariancja, ddof=0)
def slope_stat(x, y, axis=0):
    cov_xy = np.mean((x - np.mean(x)) * (y - np.mean(y)))
    var_x  = np.var(x, axis=axis, ddof=0)
    return cov_xy / var_x

a_res = bootstrap(
    data=(sample_2["X"].values, sample_2["Y"].values),
    statistic=slope_stat,
    paired=True,
    vectorized=False,
    confidence_level=conf_level,
    random_state=seed
)
a_ci = [a_res.confidence_interval.low,
        a_res.confidence_interval.high]

```

<!-- grade_id: cell-cc1ee5d5033d0923, points: 0, grade: false, solution: true -->
```python
# b
import numpy as np
from scipy.stats import bootstrap

alpha = 0.02
conf_level = 1 - alpha
seed = 25

# 6) b_ci – wyraz wolny prostej regresji
def intercept_stat(x, y, axis=0):
    a = slope_stat(x, y, axis)
    return np.mean(y) - a * np.mean(x)

b_res = bootstrap(
    data=(sample_2["X"].values, sample_2["Y"].values),
    statistic=intercept_stat,
    paired=True,
    vectorized=False,
    confidence_level=conf_level,
    random_state=seed
)
b_ci = [b_res.confidence_interval.low,
        b_res.confidence_interval.high]
```

<!-- grade_id: cell-34b28c6c1e5083f0, points: 7, grade: true, solution: false -->
```python
# Komórka testowa
assert np.all(np.isclose(mean_ci, (4.890098551749027, 5.063500727815335)))
assert np.all(np.isclose(var_ci, (0.013001391902892368, 0.05670851692597062)))
assert np.all(np.isclose(std_ci, (0.11423521149909441, 0.23829111055240196)))
assert np.all(np.isclose(r_ci, (0.9156715777126913, 0.9822743695601276)))
assert np.all(np.isclose(a_ci, (1.792948341896073, 2.3282111601384945)))
assert np.all(np.isclose(b_ci, (-5.700934784848322, -3.4292283266002834)))
```

```python

```

```python

```